```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

<div class="alert alert-info">
  <img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>
  <h1>Lab 2 – Tables and Data Manipulation</h1>
  <h2>Data 6, Fall 2025</h2>
  <p>
    In this lab, we will be talking all about <em>Tables</em>. We use tables to store all sorts of data from sports statistics to population information. If there's data you have ever been curious about, it is very likely that the Internet has a table somewhere with that data!
  </p>
  <p>
    Tables are integral to the foundation of Data Science, and we will go over how to <strong>query</strong> a table. <strong>Querying</strong> a table is, simply put, requesting information about the table. Some examples of common queries (in English, not code):
  </p>
  <ul>
    <li>How many data points are there?</li>
    <li>Which data points have a specific characteristic?</li>
    <li>What is the attribute of a specific data point?</li>
    <li>And many more!</li>
  </ul>
  <p>
    There are so many ways we can use tables to get information we need, and there are several existing libraries in Python that we can use to do this! In this course, we will be using the <code>datascience</code> library, which we've already talked about in class. This is the standard library used both in Data 6 and Data 8 at UC Berkeley (as well as when introducing data science at many other colleges). If you take Data Science classes when you enter college, you'll learn more!
  </p>
</div>


```python
# Run this cell to load all required Python libraries 
import numpy as np
from datascience import *

import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")

import warnings
warnings.simplefilter('ignore')
```

## Table Creation

Let's take a look at a table in action. Python does not have any tables by default, so we can either *create a new table from scratch* or *import a table from a file*. First, let's see how we can make our own table from scratch.

We start out with an empty `Table` -- this is the same idea as having an empty array or string. Note that `Table` is capitalized and there is nothing in the parentheses.

```python
# Run this cell to load an empty table
our_table = Table()
our_table
```

## Adding Data: `with_columns`

<div class="alert alert-success">
  <p><strong>Now, let's put some data in our table!</strong></p>
  <p>
    To do so, we use the <code>with_columns</code> method. This method requires <strong>two arguments</strong>:
  </p>
  <ol>
    <li>The name of the column as a string</li>
    <li>An array of values to put in the column</li>
  </ol>
  </div>
  <p>
    An example call to <code>with_columns</code> looks like: 
    <code>my_tbl.with_columns("My New Column", my_array)</code>, where 
    <code>my_tbl</code> is our table that we would like to add to and 
    <code>my_array</code> is a previously-defined array.
  </p>
  <p>Run the cell below to see how we can add multiple columns into <code>our_table</code>.</p>



```python
our_table = our_table.with_columns(
    "Department", make_array("Data Science", "Economics", "Political Science", "Sociology"),
    "Course Number", make_array(6, 1, 2, 121)
)
our_table
```

We need to make sure that the columns we add to the table have the same number of rows (the length of the array we pass in) as the table. Otherwise, we'll get an error.

Watch what happens if we try to add a new column that doesn't have enough data (you'll see an error!)

```python
# Just run this cell
our_table.with_columns("Too Few Rows", np.array([1, 2, 3]))
```

<div class="alert alert-primary">
<b>Question 1.1</b>: Add a new column to `our_table` called `"Number of Students"` that contains the number of students in each department. James tells you:
<ul>
<li>the <i>Data Science</i> department has 240 students,</li><li> the <i>Economics</i> department has 905 students,</li><li> the <i>Political Science</i> Department has 209 students,</li><li> and the <i>Sociology</i> department has 63 students. </li></ul>

Assign this new table to the variable `our_table_new_column`
</div>

```python
# Our table has 4 rows, so our new column needs an array with 4 items, 1 for each row
students_array = ...
our_table_new_column = ...
our_table_new_column
```

```python
grader.check("q1_1")
```

```python

```

```python
# This is our final table!
# You may use this cell to explore the table and see what you can do with it so far!
our_table_new_column
```

## Table attributes: `num_rows` and `num_columns`

We can ask for all sorts of information about the table itself:

```python
our_table_new_column.num_rows
```

```python
our_table_new_column.num_columns
```

## Accessing Data: `column`

We can also ask about the data in the table using the `column` method. As mentioned in lecture, we can pass in a `label` or an `index` to this method. 

*Note*: Recall that index into the columns of a table using **zero-indexing** -- `0` corresponds to the first column, `1` corresponds to the second, etc....

```python
# Converts a column in the table to an array
our_table_new_column.column("Department")
```

```python
# Same thing, but uses column index instead of label
our_table_new_column.column(0)
```

<div class="alert alert-primary">
  <p><strong>Question 1.2</strong></p>
  <p>
    Find the average number of students in each department by first accessing the 
    <code>"Number of Students"</code> column as an array and then taking the average. 
    Assign the average to <code>avg_num_students</code>.
  </p>
  <p><em>Note:</em> You may use any <code>np</code> functions here.</p>
</div>




```python
avg_num_students = ...
avg_num_students
```

```python
grader.check("q1_2")
```

## Loading a Table

Although creating our own tables by hand can be useful, more often than not the data we want to work with is far too large to be able to type out by hand. More commonly, we load datasets in from other sources using the `Table.read_table()` method. We can pass in a *file path* to this method and it will load that data into a table we can use in Python!

<div class="alert alert-warning">
  <h3>Background on the Data</h3>
  <p>
    The dataset that we'll use in this lab comes from the Behavioral Risk Factor Surveillance System (BRFSS), a health survey fielded by the Centers for Disease Control and Prevention (CDC). From the
    <a href="https://www.cdc.gov/brfss/index.html">BRFSS website</a>:
  </p>
  <t><b>
    <p>The Behavioral Risk Factor Surveillance System (BRFSS) is the nation’s premier system of health-related telephone surveys that collect state data about U.S. residents regarding their health-related risk behaviors, chronic health conditions, and use of preventive services.</p>
  </b>
  <b>
    <p>By collecting behavioral health risk data at the state and local level, BRFSS has become a powerful tool for targeting and building health promotion activities.</p>
  </b>
  <p>
    The dataset that you will investigate is a <strong>subset of the 2022 BRFSS Survey</strong>. We've taken all the data points corresponding to fully-completed surveys and, in our opinion, the most interesting columns. Since the entire data set is so large, we've randomly sampled 100,000 rows from the original data. While we've wrangled and cleaned the dataset you'll use in your investigation, you're welcome to investigate the original source; you can do so via the
    <a href="https://www.cdc.gov/brfss/data_documentation/index.htm">Survey Data section</a>
    of the BRFSS site.
  </p>
</div>


## Seeing a table: `show()`


<div class="alert alert-success">
The use of the <code>show()</code> method **displays** the first n rows of a table. Like <code>print()</code> this does not return a value, it just displays the value to us at the end of a cell.
</div>

```python
# Run this cell to load in the data
brfss = Table.read_table("brfss.csv")
brfss.show(5)
```

<div class="alert alert-primary">
  <p><strong>Question 2.1:</strong> Fill in the <code>num_rows_brfss</code> and <code>num_columns_brfss</code> with the number of rows and columns in the original <code>brfss</code> table, respectively.</p>
</div>


```python
num_rows_brfss = ...
num_columns_brfss = ...
print(f"Our `brfss` table has {num_rows_brfss} rows and {num_columns_brfss} columns")
```

```python
grader.check("q2_1")
```

### Investigating our Data

Now that we've loaded our data into the `brfss` table, let's take a closer look at its columns. Run the following cell to output the column names. 

```python
brfss.labels
```

Based on these column names, it looks like the data includes questions about **telecommunications**, **housing**, **demographic information**, **mental and physical health**, **alcohol and drug consumption**, and **physical exercise**. Each column in the `brfss` table corresponds to a question asked in the official BRFSS Survey.

## Data Manipulation

Now that we have a solid understanding of the basic table methods from the `datascience` library, let's put our skills to use! Even with a few tools, we are already able to arrive at powerful realizations about real world data.

<div class="alert alert-primary">
  <p><strong>Question 2.2:</strong> Assign <code>num_alabamaa_rows</code> to the the number of times the name <b>Alabama</b> appeared in the <code>brfss</code> table.</p>

</div>

*Hint*: Use the table methods you've learned above!


```python
alabama_tbl = ...
num_alabama_rows = ...
num_alabama_rows
```

```python
grader.check("q2_2")
```

```python
missing_binge_drinks = ...
missing_binge_drinks
```

<div class="alert alert-primary">
  <p><strong>Question 2.4 (<i>Discussion</i>):</strong> Say we wanted to find the average of one of the columns from our original table. How does the inclusion of <code>-1</code> values <i>affect</i> this average? If we removed all the negative values, how would the average change?
</p>
<p>Then, once you've answered, run the following cell to confirm your understanding.
</p></div>


_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
brfss_no_negatives_children = brfss.where("Children", are.not_equal_to(-1))
print(f"With negatives: {np.average(brfss.column('Children'))}")
print(f"Without negatives: {np.average(brfss_no_negatives_children.column('Children'))}")
```

<div class="alert alert-primary">
  <p><strong>Question 2.5:</strong>
  As an <b><i>extra challenge</i></b>, see if you can output Alabama's average number of children for all respondents who gave an answer (i.e. do not have a value of -1) <b>without using the <code>np.average</code> function</b>. You should use the table <code>brfss_no_negatives_children</code> that's defined in the cell above.

</div>
<i>Note</i>: I've provided some starter code. Feel free to use it or try another approach!


```python
alabama_tbl = ...
children_total = ...
average_num_children = ...
average_num_children
```

```python
grader.check("q2_5")
```

<div class="alert alert-primary">
  <p><strong>Question 2.6:</strong> Using the <code>no_missing_income</code> table we've provided for you, determine the <strong>average income</strong> for respondents who:</p>
  <ol>
    <li>Have a health insurance plan ("Yes" in <code>Health Plan</code> column)</li>
    <li>Do not have a health insurance plan ("No" in <code>Health Plan</code> column)</li>
    <li>Refused to answer this question ("Declined to Answer" in <code>Health Plan</code> column)</li>
  </ol>
  <p>You may use the starter code provided for you, but are not required to.</p>
</div>


```python
no_missing_income = brfss.where("Income Upper", are.not_equal_to(-1))

health_plan = ...
no_health_plan = ...
declined = ...

average_plan = ...
average_no_plan = ...
average_declined = ...

print(f"Respondents with a Health Insurance Plan made: \t\t${round(average_plan, 2)}")
print(f"Respondents without a Health Insurance Plan made: \t${round(average_no_plan, 2)}")
print(f"Respondents who refused to answer made: \t\t${round(average_declined, 2)}")
```

```python
grader.check("q2_6")
```

<div class="alert alert-success">
With just over three weeks of data science under your belts, you're already able to uncover underlying <b>patterns and trends</b> within real world data. In this case, we've found that those with health insurance plans make, on average, almost 9,000 more dollars a year than those without health insurance plans.

What's valuable, too, is the information we gain when a population is <i>missing</i> or, in this case, <i>declines to participate</i>. Notice that the average salary of respondents who refused to answer was around 4,000 dollars lower than those *without* a Health Insurance Plan.
</div>

## Done! 👉✅

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: 1)
```
>>> our_table_new_column.labels == ('Department', 'Course Number', 'Number of Students')
True
```
```
>>> all(our_table_new_column.column('Number of Students') == make_array(240, 905, 209, 63))
True
```


Tests for q1_2 (points: 1)
```
>>> avg_num_students == 354.25
True
```


Tests for q2_1 (points: 1)
```
>>> num_rows_brfss == 100000
True
```
```
>>> num_columns_brfss == 23
True
```


Tests for q2_2 (points: 1)
```
>>> num_alabama_rows == 1352
True
```


Tests for q2_3 (points: 1)
```
>>> missing_binge_drinks.num_rows == 51042
True
```


Tests for q2_5 (points: 1)
```
>>> np.isclose(average_num_children, 0.5443698732289336)
True
```


Tests for q2_6 (points: 1)
```
>>> np.isclose(average_plan, 43832.35496992514)
True
```
```
>>> np.isclose(average_no_plan, 34940.17647058824)
True
```
```
>>> np.isclose(average_declined, 30963.912280701756)
True
```
